# BLU10 - Exercises Notebook

```python
# Import all the necessary dependencies
import os
import numpy as np
import scipy as sp
import pandas as pd

import scipy.sparse

from mlxtend.frequent_patterns import apriori
import hashlib # for grading purposes
```

## Q0: Create the ratings matrix (ungraded)

```python
path = os.path.join('data', 'ml-latest-small', 'ratings.csv')
data = pd.read_csv(path)
# Shuffle Data
data = data.sample(10493, random_state=200)
data.head()
```

Output:
```
userId  movieId  rating   timestamp
58111     423     1206     5.0  1353691236
94054     624     3268     1.0  1028111170
97308     652    26843     5.0  1440269953
55435     401      924     5.0   977458816
22437     157     5378     2.5  1323618006
```

So, we'll ignore the _timestamp_ and use the _rating_ column as our only source of information for our recommender system. Remember that if you had data of other types of interactions between users and the movies, you could create a metric that incorporates all of that information in the ratings matrix (maybe by averaging them).

Keep the following ratings matrix schema in your mind while developing non-personalized systems. These systems rely heavily on the ratings matrix, so maybe also write it on a piece of paper to remember it better!

<img align="left" width="413" height="239" src="./media/ratings_matrix3.png">


In the following exercise you will build the ratings matrix with users as rows and products as columns.

Tip: you can use the pandas' _pivot_ function or flex your numpy muscles with the _genfromtxt_ function (it is good for your health!).

<!-- grade_id: cell-8a6b8634bc3ef2ff, points: 0, grade: false, solution: true -->
```python
def make_ratings(original_data: pd.DataFrame):
    """
    Parameters
        original_data - the original data with ratings per user and product.
        
    Returns:
        R - (numpy.ndarray) Ratings matrix with the userId, movieId and rating
        hint: don't forget to put zeros on places where you do not have ratings
    
    Extra Hint: Your input is a pandas DataFrame but you want to output an array (use .to_numpy)!
    """
    # YOUR CODE HERE
    
    return data.pivot(index='userId', columns='movieId',values='rating').fillna(0).to_numpy()
    

R = make_ratings(data)
R

```

Output:
```
array([[0., 0., 0., ..., 0., 0., 0.],
       [0., 0., 0., ..., 0., 0., 0.],
       [0., 0., 0., ..., 0., 0., 0.],
       ...,
       [0., 0., 0., ..., 0., 0., 0.],
       [0., 0., 0., ..., 0., 0., 0.],
       [0., 0., 0., ..., 0., 0., 0.]])
```

```python
f"We have {R.shape[0]} user and {R.shape[1]} items."
```

Output:
```
'We have 657 user and 3628 items.'
```

```python
expected_hash = '0825c15053e635376af0a569e8f37cfaef0e1dfce37ae6878517e14e061f13c4'
assert hashlib.sha256(str(R.shape).encode()).hexdigest() == expected_hash

expected_hash_1 = '8ab31b5afaea56114427e1f01b81d001b079a0f59539f6db3f099816ca794055'
assert hashlib.sha256(str(R[0].sum()).encode()).hexdigest() == expected_hash_1

expected_hash_2 = 'b5967724d1225caa9c6af28a9b333a29e6d5c11a24e9d381acf5c3377524b776'
assert hashlib.sha256(str(R[:,0].sum()).encode()).hexdigest() == expected_hash_2


```

## Q1: Convert the Ratings Matrix to a Sparse Representation

In this exercise, convert the ratings matrix to a sparse row representation. 

Hint: Remember what we have done with scipy library!

<!-- grade_id: cell-c0c272ba37054416, points: 0, grade: false, solution: true -->
```python
from scipy.sparse import csr_matrix

def get_csr(orig_matrix):
    """
    Parameters
        orig_matrix - The original Ratings Matrix.
    
    Returns
        H_ - The Compressed Sparse Row Matrix
    """
    # YOUR CODE HERE
    return csr_matrix(orig_matrix)
    
sparse_mat = get_csr(R)
```

<!-- grade_id: cell-35b3a4b1f0a76185, points: 2, grade: true, solution: false -->
```python
expected_hash = 'e33275c9e0741880dc0334b76fb8cb17e27020dec75dda07ab221dbb97277d30'
assert hashlib.sha256(str(sparse_mat).encode()).hexdigest() == expected_hash
```

## Q2: What is the density score of this matrix?

In this exercise, let's understand the density score (this is, the fraction of rows that are non zero in the original rating matrix).

Calculate the get_density_score function below:

<!-- grade_id: cell-49adf0f3c0101526, points: 0, grade: false, solution: true -->
```python
def get_density_score(orig_matrix):
    """
    Parameters
        orig_matrix - Ratings Matrix
        
    Returns:
        dense_score - (float) Density Score of matrix Orig Matrix. 
    """
    # YOUR CODE HERE
    # YOUR CODE HERE
    return orig_matrix[orig_matrix.nonzero()].size / orig_matrix.size
    
dense_score = get_density_score(R)
dense_score
```

Output:
```
0.004402172180184897
```

<!-- grade_id: cell-0792440f45999d40, points: 2, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(dense_score,0.004,3)
```

```python
f"The Density Score is {dense_score}."
```

Output:
```
'The Density Score is 0.004402172180184897.'
```

## Q3: Popular Items - What are the Top-3 Most Rated items?
More ratings give us the current trends but not necessarily the best suggestions - but let's check Items that have more ratings given.

In this exercise you will have to retrieve the indexes of the products so you may need to recreate the ratings matrix as a dataframe or come up with another creative solution!

Return the product IDs from the product with most ratings to the lowest.

**Hint: To get the ID's it's easier if you work with the original data frame other than the rating matrix!**
**Remember the pandas pivot method! ;)**

<!-- grade_id: cell-9ce05dd01331575e, points: 0, grade: false, solution: true -->
```python
def top3items(original_df, n):
    """
    Parameters
        original_df - Original Data Frame with ratings
        n - Number of Top-n items to retrieve
        
    Returns
        top_ids - (list) list of product ids of  
        the top-n most rated items
    """
    # YOUR CODE HERE
    # YOUR CODE HERE
    
    return list(original_df.pivot(index='userId',columns='movieId',values='rating').count(axis=0).sort_values(ascending= False).index[:n])

most_wanted = top3items(data, 3)

most_wanted
```

Output:
```
[318, 356, 296]
```

<!-- grade_id: cell-ae4661f44fc09de8, points: 3, grade: true, solution: false -->
```python
expected_hash = '3f54e8a66e4924557721acdcb1d40e4399376f780ce6d2b11f48931fc5e1376f'
assert hashlib.sha256(str(most_wanted).encode()).hexdigest() == expected_hash
```

## Q4: Influencers - What are the Top-5 Most Active Users?

Now let's do the following for users! Return the top 5 most active users - this is, the ones that have rated most products. 

Return the user IDs from the users with most ratings to the lowest ones (but only the top 5).

<!-- grade_id: cell-1559bfc72b8eaf4c, points: 0, grade: false, solution: true -->
```python
def get_influencers(original_df, n):
    """
    Parameters
        original_df - Original DataFrame with ratings
        n - number of top-n most active users
        
    Returns
        influencers - (list) list of 
        ids of the top-n most active users
    """
    
    # YOUR CODE HERE
    influencers=list(original_df.pivot(index='movieId',columns='userId',values='rating').count(axis=0).sort_values(ascending= False).index[:5])
    
    return influencers

influencers = get_influencers(data, 5)

influencers
```

Output:
```
[547, 564, 15, 624, 73]
```

<!-- grade_id: cell-2cd4b77837e74779, points: 4, grade: true, solution: false -->
```python
expected_hash = '2cbf47fd3da8e22be1bff24eb3e2ad0b99b8992d870e7a5ea4f21880588accc6'
assert hashlib.sha256(str(influencers).encode()).hexdigest() == expected_hash
```

## Q5: Elite - What are the Top-7 Better Rated Items (On Average)?

Since this can be biased by a low number of ratings, we need items to have more than 10 ratings. Use average to obtain the ids of the top average rated products. Return the product ID from the highest rated to the lowest rated item.

Hint: In this exercise and to filter the movies by rating, it may be easier to use the original data and then reconstruct the ratings matrix!
<br>
Hint 2: Don't forget that we are asking for the **top** rated items so you have to sort your average ratings in some way!

<!-- grade_id: cell-62b63f8f75bd5d58, points: 0, grade: false, solution: true -->
```python
def elite(original_data, n, k):
    """
    Parameters
        original_data - The original dataframe with ratings.
        n - Top-n items
        k - Mininum number of ratings
        
    Returns
        best_items - (list) list of ids of top-n best mean rated items.
        Your indices should refer only to items with more than k ratings (subset of original matrix).
    """
    
    R_ = original_data.pivot(index='userId',columns='movieId',values='rating').fillna(0)
    
    count_ratings = np.greater(R_, 0).sum(axis=0)
    
    R_[R_ == 0] = np.NaN
    mean_ratings = np.nanmean(R_, axis=0)
    
   
    df = pd.DataFrame([pd.Series(count_ratings),pd.Series(mean_ratings, index=count_ratings.index)]).T
    best_items = list(df[df[0]>k].sort_values(by=[1], ascending= False).index[:n])
    
    return best_items
    
best_items = elite(data, 7, 10)

best_items

```

Output:
```
[527, 2959, 50, 318, 1193, 1219, 2571]
```

<!-- grade_id: cell-3e673279793bc5d4, points: 5, grade: true, solution: false -->
```python
expected_hash = '0ed83bd0067fb63b2c591d0039829a3a9bc6887a7bb8f076bc15f6e6944e9216'
assert hashlib.sha256(str(best_items).encode()).hexdigest() == expected_hash
```

## Q6: Apriori - What are the 5 most common 2-piece itemsets?
We define "common itemsets" as at least 2 different items that are usually rated together at least by 0.5% of the population (erheeem support!).
Show your results sorted by support in descending way.

Hint: Check the mlxtend documentation for help: http://rasbt.github.io/mlxtend/user_guide/frequent_patterns/apriori/

<!-- grade_id: cell-dff47ecd831f9ca3, points: 0, grade: false, solution: true -->
```python
def getBundlesSolution(original_data, n=None, min_support=None, top=None):
    """
    Parameters
        original_data: Original data frame with ratings.
        n: Number of items in commonset
        min_support: Minimum percentage of users that contains the itemset
        top: Number of most common itemsets
        
    Return
        df: the return dataframe should have two columns ["support", "itemsets"],
            with the support percentage and the itemsets.
    """
    R = original_data.pivot(index='userId',columns='movieId',values='rating').fillna(0)
    R_ = pd.DataFrame(R > 0)
    
    apriori_mt = apriori(R_, min_support)
    apriori_mt_n = apriori_mt[apriori_mt.itemsets.map(len)==n]
    return apriori_mt_n.sort_values(by='support', ascending = False)[:top]
 

df = getBundlesSolution(data, n=2, min_support=0.005, top=5)
df
```

Output:
```
support    itemsets
834  0.010654  (268, 189)
846  0.010654  (229, 279)
823  0.009132  (172, 229)
847  0.009132  (229, 342)
813  0.009132  (177, 155)
```

<!-- grade_id: cell-404e9c98f711743d, points: 4, grade: true, solution: false -->
```python
expected_hash = 'f7441550a0ca5274581d023417c99540e3a8a4cca68824a87cbe6d95c07742ea'
assert hashlib.sha256(str(df.shape).encode()).hexdigest() == expected_hash

np.testing.assert_almost_equal(df.iloc[0,0],0.011,3)

np.testing.assert_almost_equal(df.iloc[4,0],0.009,3)
```

```python

```