# Basic Workflow

```python
# Always have your imports at the top
import random
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import MinMaxScaler
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score
from sklearn.base import TransformerMixin
from sklearn.model_selection import train_test_split
from sklearn.impute import SimpleImputer

from hashlib import sha1 # just for grading purposes
import json # just for grading purposes

from utils import get_dataset, workflow_steps, data_analysis_steps

def _hash(obj, salt='none'):
    if type(obj) is not str:
        obj = json.dumps(obj)
    to_encode = obj + salt
    return sha1(to_encode.encode()).hexdigest()

X, y = get_dataset()  # preloaded dataset
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33)
```

## Exercise 1: ~The Larch~ Workflow steps

What are the basic workflow steps?

<img src="media/the_larch.jpg" width="300" />

<p style="font-size: 9px; text-align:center">A larch</p>

You probably know them already, but we want you to really internalize them. We've given you a list of steps `workflow_steps`, but it appears that, not only does it have to many steps, some are _probably_ wrong, as well.

Select the correct ones and reorder them!

```python
print("Workflow steps:")
for i in range(len(workflow_steps)):
    print(i+1, ': ', workflow_steps[i])
```

Output:
```
Workflow steps:
1 :  Data analysis and preparation
2 :  Google Hackathon solutions
3 :  Spam
4 :  Iterate
5 :  Establish a Baseline
6 :  Get the data
7 :  Watch Netflix
8 :  Train model
9 :  Evaluate results
10 :  Increase complexity
```

<!-- grade_id: cell-6c8218e1615228f0, points: 0, grade: false, solution: true -->
```python
# Exercise 1.1. Filter and sort the names of the steps in the workflow_steps list
workflow_steps_answer = ['Get the data', 'Data analysis and preparation', 'Train model', 'Evaluate results', 'Iterate']
workflow_steps_answer
```

Output:
```
['Get the data',
 'Data analysis and preparation',
 'Train model',
 'Evaluate results',
 'Iterate']
```

<!-- grade_id: cell-d15af97ae329a075, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
assert _hash([step.lower() for step in workflow_steps_answer], 'salt0') == '701e2306da9bfde36382bdb6feb80a354916ebf4'
### END TESTS
```

There are way too many substeps in the Data Analysis and Preparation step to group them all under a single category. We've given you another list of steps: `data_analysis_steps`.

Aside from being shuffled, it should be fine but keep an eye out. You never know what to expect...

```python
print("Data Analysis and Preparation steps:")
for i in range(len(data_analysis_steps)):
    print(i+1, ': ', data_analysis_steps[i])
```

Output:
```
Data Analysis and Preparation steps:
1 :  Feature selection
2 :  Dealing with data problems
3 :  Data analysis
4 :  Feature engineering
5 :  Spanish Inquisition
```

<!-- grade_id: cell-1e47f7838cbcf13d, points: 0, grade: false, solution: true -->
```python
# Exercise 1.2. Filter and sort the names of the steps in the data_analysis_steps list
data_analysis_steps_answer = ['Data analysis','Dealing with data problems', 'Feature engineering', 'Feature selection']

```

<!-- grade_id: cell-eb6774d7694e5ada, points: 2, grade: true, solution: false -->
```python
### BEGIN TESTS
assert _hash([step.lower() for step in data_analysis_steps_answer], 'salt0') == '658ab90eff4a0cea2bfb51cc89c8db5b4121fa86'
### END TESTS
```

<p style="text-align:center">That's right! 

<p style="text-align:center"><b>Nobody</b> expects the Spanish Inquisition!

<img src="media/spanish_inquisition.gif" width="400" />

## Exercise 2: Specific workflow questions

Here are some more specific questions about individual workflow steps.

<!-- grade_id: cell-53dadb4a6c1cb987, points: 0, grade: false, solution: true -->
```python
# Exercise 2.1. True or False, you should split your data in a training and test set
split_training_test = True

# Exercise 2.2. True or False, Scikit Pipelines are only useful in production environments
scikit_pipelines_useful = False

# Exercise 2.3. True or False, you should try to make a complex baseline, so you just have 
#               to make simple improvements on it, later on.
baseline_complex = False

# Exercise 2.4. (optional) True or False, is Brian the Messiah?
# is_brian_the_messiah = ...

```

<!-- grade_id: cell-f1cfeb62360089e3, points: 6, grade: true, solution: false -->
```python
### BEGIN TESTS
assert _hash(split_training_test, 'salt1') == '569b45c42b5c7b490c92692b911af35f575c8a06'
assert _hash(scikit_pipelines_useful, 'salt2') == 'ef07576cc7d3bcb2cf29e1a772aec2aad7f59158'
assert _hash(baseline_complex, 'salt3') == 'f24a294afb4a09f7f9df9ee13eb18e7d341c439d'
### END TESTS
```

<img src="media/monty_python_messiah.jpg" width=500>

## Scikit Pipelines

We've already loaded and splitted a dataset for the following exercises. They're stored in the `X_train`, `X_test`, `y_train` and `y_test` variables.

In a perfect world, where you have all your data clean and ready-to-go, you can create your pipeline with just Scikit-learn's Transformers. However, in the real world, that's not the case, and you'll need to create custom Transformers to get the job done. Take a look at the data set, what do you see?

```python
# Do some data analysis here
X_train.describe()
X_train.info()
X_train.head()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 67 entries, 18 to 46
Data columns (total 24 columns):
 #   Column  Non-Null Count  Dtype  
---  ------  --------------  -----  
 0   17      67 non-null     float64
 1   leg_1   67 non-null     object 
 2   7       67 non-null     float64
 3   arm_2   67 non-null     object 
 4   2       67 non-null     float64
 5   11      67 non-null     float64
 6   4       61 non-null     float64
 7   16      67 non-null     float64
 8   0       62 non-null     float64
 9   14      57 non-null     float64
 10  5       59 non-null     float64
 11  3       67 non-null     float64
 12  leg_3   67 non-null     object 
 13  13      67 non-null     float64
 14  12      60 non-null     float64
 15  10      67 non-null     float64
 16  15      67 non-null     float64
 17  9       67 non-null     float64
 18  19      67 non-null     float64
 19  18      67 non-null     float64
 20  8       67 non-null     float64
 21  6       67 non-null     float64
 22  arm
```

While crunching your data, you probably found two issues:

1. There are 4 columns whose name starts with either `arm` or `leg` which are all filled with gibberish
2. There are some values missing in some columns

So, first things first, let's get rid of those columns through a Custom Transformer, so we can plug it in a Scikit Pipeline after.

## Exercise 3: Custom Transformer

<!-- grade_id: cell-e9663ebd68ad68f8, points: 0, grade: false, solution: true -->
```python
# Create a pipeline step called RemoveEvilColumns that removes any
# column whose name starts with the string 'evil'

class RemoveEvilColumns(TransformerMixin):
    
    def transform(self, X, *_):
        columns= X.columns
        list_evil = []
        for col in columns:
            if col[:5] == 'evil':
                list_evil.append(col)
                
        new = X[list_evil].copy()
        
        return new


    def fit(self, *_):
        return self
   
```

<!-- grade_id: cell-5c8634dcc19dd7e1, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert _hash(sorted(RemoveLimbs().fit_transform(X).columns), 'salt5') == '71443dfc3077d773d4c74e958dadf91dc2cc148a'
assert _hash(list(map(lambda col: col.startswith('arm') or col.startswith('leg'), RemoveLimbs().fit_transform(X_train).columns)), 'salt6') == 'ce45cf3759d2210f2d1315f1673b18f34e3ac711'
### END TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-105-82c9da05fdd6> in <module>
      1 ### BEGIN TESTS
----> 2 assert _hash(sorted(RemoveLimbs().fit_transform(X).columns), 'salt5') == '71443dfc3077d773d4c74e958dadf91dc2cc148a'
      3 assert _hash(list(map(lambda col: col.startswith('arm') or col.startswith('leg'), RemoveLimbs().fit_transform(X_train).columns)), 'salt6') == 'ce45cf3759d2210f2d1315f1673b18f34e3ac711'
      4 ### END TESTS

NameError: name 'RemoveLimbs' is not defined
```

<img src="media/monty_python_black_knight.gif" width=500>

Now that we have our Custom Transformer in place, we can design our pipeline. For the sake of the exercise, you'll want to create a pipeline with the following steps:

1. Removes limbs columns
2. Imputes missing values with the mean
3. Has a Random Forest Classifier as the last step

You may use `make_pipeline` to create your pipeline with as many steps as you want as long as the first two are the Custom Transformer you developed previously, a `SimpleImputer` as the second step, and a `RandomForestClassifier` as the last step.

## Exercise 4: Scikit Pipelines

<!-- grade_id: cell-efabb79e960ce10b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6814224bb0f1332f, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert _hash(pipeline.steps[0][0], 'salt7') == '471b02068ac2c4f479c2e9f85f4b3dc2179bb841'
assert _hash(pipeline.steps[1][0], 'salt8') == 'ca83eaea1a7e243fa5574cfa6f52831166ee0f32'
assert _hash(pipeline.steps[-1][0], 'salt9') == '0d66ba4309ad4939673169e74f87088dcadd510b'
### END TESTS
```

Does it work? Let's check it out on our dataset!

```python
pipeline.fit(X_train, y_train)
```

```python
y_pred = pipeline.predict(X_test)
accuracy_score(y_test, y_pred)
```

That's it for this Exercise Notebook! It doesn't get much cleaner than this, does it?

You can still practice around with pipelines, maybe add a few more steps. See how you can adapt your pipeline and how it affects the predictions.

Can you see how Scikit-learn's Pipelines might save time? Can you imagine how useful that would be in stressful situations (like, *for example*, an Hackathon)?